Before you submit this notebook, make sure everything runs as expected in the local test cases. 
Please, paste the solution to the designed cell and do not change anything else.

Also, please, leave your first and last names below

```python
FirstName = ""
LastName = ""
```

---

```python
import sys

import numpy as np
import unittest
import time

import collections
import pickle
import io
```

<!-- grade_id: cell-ac8fc52d8a39ccb4, points: 0, grade: false, solution: true -->
```python
import numpy as np


class LossAndDerivatives:
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Test 0: Initialization (0.01 points)

<!-- grade_id: cell-3473b7b6ffd64d07, points: 0.01, grade: true, solution: false -->
```python
# do not change this cell
loss_and_derivatives = LossAndDerivatives


```

### Test 1: MSE derivative (0.24 points)

<!-- grade_id: cell-e3503c286039ec55, points: 0.24, grade: true, solution: false -->
```python
mse_derivative = LossAndDerivatives.mse_derivative(X_ref, y_ref, w_hat)
assert np.allclose(mse_derivative, ref_dict['mse_derivative'], atol=1e-4)
```

### Test 2: MAE derivative (0.25 points)

<!-- grade_id: cell-e2c4124a6f815118, points: 0.25, grade: true, solution: false -->
```python
mae_derivative = LossAndDerivatives.mae_derivative(X_ref, y_ref, w_hat)
assert np.allclose(mae_derivative, ref_dict['mae_derivative'], atol=1e-4)
```

### Test 3: L1 reg derivative (0.15 points)

<!-- grade_id: cell-69473387a23d8dff, points: 0.15, grade: true, solution: false -->
```python
l2_reg_derivative = LossAndDerivatives.l2_reg_derivative(w_hat)
assert np.allclose(l2_reg_derivative, ref_dict['l2_reg_derivative'], atol=1e-4)
```

### Test 4: L1 reg derivative (0.15 points)

<!-- grade_id: cell-3460671ba79fd04a, points: 0.15, grade: true, solution: false -->
```python
l1_reg_derivative = LossAndDerivatives.l1_reg_derivative(w_hat)
assert np.allclose(l1_reg_derivative, ref_dict['l1_reg_derivative'], atol=1e-4)
```

### Test 5: MSE (0.05 points)

<!-- grade_id: cell-5d56bb4222ee1e7c, points: 0.05, grade: true, solution: false -->
```python
mse = LossAndDerivatives.mse(X_ref, y_ref, w_hat)
assert np.allclose(mse, ref_dict['mse'], atol=1e-4)
```

### Test 6: MAE (0.05 points)

<!-- grade_id: cell-8155e619427ac916, points: 0.05, grade: true, solution: false -->
```python
mae = LossAndDerivatives.mae(X_ref, y_ref, w_hat)
assert np.allclose(mae, ref_dict['mae'], atol=1e-4)
```

### Test 7: L2 reg (0.05 points)

<!-- grade_id: cell-c7ccfce79ecbadc3, points: 0.05, grade: true, solution: false -->
```python
l2_reg = LossAndDerivatives.l2_reg(w_hat)
assert np.allclose(l2_reg, ref_dict['l2_reg'], atol=1e-4)
```

### Test 8: L1 reg (0.05 points)

<!-- grade_id: cell-1e1aa415f6d381f7, points: 0.05, grade: true, solution: false -->
```python
l1_reg = LossAndDerivatives.l1_reg(w_hat)
assert np.allclose(l1_reg, ref_dict['l1_reg'], atol=1e-4)   
```