Before you submit this notebook, make sure everything runs as expected in the local test cases. 
Please, paste the solution to the designed cell and do not change anything else.

Also, please, leave your first and last names below

```python
FirstName = ""
LastName = ""
```

---

```python
import numpy as np
from sklearn.base import BaseEstimator
from sklearn.preprocessing import OneHotEncoder
```

<!-- grade_id: cell-ac8fc52d8a39ccb4, points: 0, grade: false, solution: true -->
```python

def entropy(y):  
    """
    Computes entropy of the provided distribution. Use log(value + eps) for numerical stability
    
    Parameters
    ----------
    y : np.array of type float with shape (n_objects, n_classes)
        One-hot representation of class labels for corresponding subset
    
    Returns
    -------
    float
        Entropy of the provided subset
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Test 0: Initialization (0.01 points)

<!-- grade_id: cell-3473b7b6ffd64d07, points: 0.01, grade: true, solution: false -->
```python
# do not change this cell
import numpy as np
import unittest
import sys
import io

from sklearn.datasets import fetch_california_housing, load_digits
from sklearn.metrics import accuracy_score, mean_squared_error
from sklearn.model_selection import train_test_split

digits_data = load_digits(n_class=2).data
digits_target = load_digits(n_class=2).target[:, None]


```

### Test 1: Make splits loops (0.1 points)

<!-- grade_id: cell-e3503c286039ec55, points: 0.09, grade: true, solution: false -->
```python
X = np.ones((4, 5), dtype=float) * np.arange(4)[:, None]
y = np.arange(4)[:, None] + np.asarray([0.2, -0.3, 0.1, 0.4])[:, None]
class_estimator = DecisionTree(max_depth=5, criterion_name='gini')

(X_l, y_l), (X_r, y_r) = class_estimator.make_split(1, 1., X, y)

flag_X = np.array_equal(X[:1], X_l) and np.array_equal(X[1:], X_r) 
flag_y = np.array_equal(y[:1], y_l) and np.array_equal(y[1:], y_r)
assert flag_X and flag_y
```

### Test 2: Gini accuracy (0.2 points)

<!-- grade_id: cell-e2c4124a6f815118, points: 0.2, grade: true, solution: false -->
```python
class_estimator = DecisionTree(max_depth=5, criterion_name='gini')
class_estimator.fit(X_train, y_train)
ans = class_estimator.predict(X_test)
accuracy_gini = accuracy_score(y_test, ans)
assert 0.96 < accuracy_gini
```

### Test 3: Entropy accuracy (0.2 points)

<!-- grade_id: cell-69473387a23d8dff, points: 0.2, grade: true, solution: false -->
```python
class_estimator = DecisionTree(max_depth=5, criterion_name='entropy')
class_estimator.fit(X_train, y_train)
ans = class_estimator.predict(X_test)
accuracy = accuracy_score(y_test, ans)
assert 0.96 < accuracy
```

### Test 4: Entropy probabilities (0.2 points)

<!-- grade_id: cell-e5f59af66e6c111b, points: 0.2, grade: true, solution: false -->
```python
class_estimator = DecisionTree(max_depth=10, criterion_name='entropy')
class_estimator.fit(X_train, y_train)
ans = class_estimator.predict(X_test)
reference = np.array([0.48611111, 0.51388889])
assert np.abs(np.sum(class_estimator.predict_proba(X_test).mean(axis=0) - reference)) < 1e-6
```

### Test 5: MSE mad_median (0.15 points)

<!-- grade_id: cell-1a9c1e3609ed9aab, points: 0.15, grade: true, solution: false -->
```python
housing = fetch_california_housing()
random_indices = np.random.choice(np.arange(len(housing.data)), 500)

regr_data = housing.data[random_indices]
regr_target = housing.target[random_indices, None]
RX_train, RX_test, Ry_train, Ry_test = train_test_split(regr_data, regr_target, test_size=0.2, random_state=42)

regressor = DecisionTree(max_depth=8, criterion_name='mad_median')
regressor.fit(RX_train, Ry_train)
predictions_mad = regressor.predict(RX_test)
mse = mean_squared_error(Ry_test, predictions_mad)
assert 0.4 < mse < 2
```

### Test 6: MSE Variance (0.15 points)

<!-- grade_id: cell-1ddb0377b6c68deb, points: 0.15, grade: true, solution: false -->
```python
housing = fetch_california_housing()
random_indices = np.random.choice(np.arange(len(housing.data)), 500)

regr_data = housing.data[random_indices]
regr_target = housing.target[random_indices, None]
RX_train, RX_test, Ry_train, Ry_test = train_test_split(regr_data, regr_target, test_size=0.2, random_state=42)

regressor = DecisionTree(max_depth=8, criterion_name='variance')
regressor.fit(RX_train, Ry_train)
predictions_mad = regressor.predict(RX_test)
mse = mean_squared_error(Ry_test, predictions_mad)
assert 0.5 < mse < 1.8
```